# Introduction to Python and Natural Language Technologies

__Laboratory 02, Type system and built-in types__

__Feb 18, 2021__


We will use the wikipedia module for some exercises. It can be installed by running pip from Jupyter:

The Python Standard Library has solutions for many of these tasks but you should solve them without relying on the standard library.

```python
! pip install wikipedia
```

Output:
```
Requirement already satisfied: wikipedia in g:\software\lib\site-packages (1.4.0)
Requirement already satisfied: beautifulsoup4 in g:\software\lib\site-packages (from wikipedia) (4.9.3)
Requirement already satisfied: requests<3.0.0,>=2.0.0 in g:\software\lib\site-packages (from wikipedia) (2.24.0)
Requirement already satisfied: soupsieve>1.2; python_version >= "3.0" in g:\software\lib\site-packages (from beautifulsoup4->wikipedia) (2.0.1)
Requirement already satisfied: certifi>=2017.4.17 in g:\software\lib\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (2020.6.20)
Requirement already satisfied: idna<3,>=2.5 in g:\software\lib\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (2.10)
Requirement already satisfied: urllib3!=1.25.0,!=1.25.1,<1.26,>=1.21.1 in g:\software\lib\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (1.25.11)
Requirement already satisfied: chardet<4,>=3.0.2 in g:\software\lib\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (3.0.4)
```

# 1. Define a function that takes a sequence as its input and returns whether the sequence is symmetric. A sequence is symmetric if it is equal to its reverse.

<!-- grade_id: cell-9063e2fba9c7157d, points: 0, grade: false, solution: true -->
```python
'''
l=[1,2,3]
l1=[1,2,3]

l==l1 # True

l=[1,2,3]
old_id=id(l)
l.reverse()
old_id == id(l) #True

'''

def is_symmetric(arg):
    revrspos=len(arg)-1
    for pos in range(0,len(arg)):
        if arg[pos] != arg[revrspos]:
            return False
        revrspos=revrspos-1
    return True
```

<!-- grade_id: cell-58b3ad7e361196e7, points: 0, grade: true, solution: false -->
```python
assert is_symmetric([1]) == True
assert is_symmetric([None, None]) == True
assert is_symmetric([]) == True
assert is_symmetric([1, 2, 3, 1]) == False
assert is_symmetric([1, "foo", "bar", "foo", 1]) == True
assert is_symmetric("abcba") == True
assert is_symmetric(list(range(100)) + list(range(99, -1, -1))) == True
assert is_symmetric(list(range(100)) + list(range(100, 0, -1))) == False
assert is_symmetric([1, "1"]) == False
```

# 2. Define a function that takes a sequence and counts the frequency of each element in the list. The function should return a dictionary of the counts.

Do not use the `collections` module.

<!-- grade_id: cell-d36bfd560f1b29a3, points: 0, grade: false, solution: true -->
```python
'''
access the value associated with a key in dictionary (i.e dictionary variable is d):
d.[key]
'''

'''
l=[1, 2, 0, 1, 1]
uniq= set(l)
len(uniq)
uniq
for i in uniq: # to access an element in a set
    print (i)

# postion of elements doesn't matter in a dictionary:
d1= {1: 3, 2: 1, 0: 1}
d2= {2: 1, 1: 3, 0: 1}
print(d1==d2) #True
'''

def count_elements(arg):
    d={}
    uniq=set(arg) # get the distinctive elements in the arg (arg can be a list, a Tuple etc)
    for i in uniq:
        d[i]=0
    #print(d)
    for x in range (0, len(arg)):
        d[arg[x]]=d[arg[x]]+1
    return d
#print(count_elements([]))

#count_elements([1, 2, 0, 1, 1])

```

<!-- grade_id: cell-5175730489a10307, points: 0, grade: true, solution: false -->
```python
# test for list
assert count_elements([1, 2, 0, 1, 1]) == {1: 3, 2: 1, 0: 1}
# test for tuple
assert count_elements((1, 2, 0, 1, 1)) == {1: 3, 2: 1, 0: 1}
# test empty list
assert count_elements([]) == {}
# test mixed types
assert count_elements([1, "1", 1]) == {1: 2, "1": 1}
```

# 3. Define a function that takes a list of lists and decided if it is a matrix. A list of lists is a matrix if each nested list is the same length.

<!-- grade_id: cell-90b341f4b426d33d, points: 0, grade: false, solution: true -->
```python
def is_matrix(arg):
    if arg==[]: # if the list is empty
        return True
    len1=len(arg[0])
    for x in range (0, len(arg)):
        if len(arg[x]) != len1:
            return False
    return True

```

<!-- grade_id: cell-218d090c410748b5, points: 0, grade: true, solution: false -->
```python
assert is_matrix([]) == True
assert is_matrix([[], []]) == True
assert is_matrix([[1, 2], [2, 1]]) == True
assert is_matrix([[1], [2, 1]]) == False
assert is_matrix([[1, 2], [2, 1], [3, 4]]) == True
assert is_matrix([[1, 2], [2, 1], [3, 4, 5]]) == False
```

# 4. Define a function that takes a string as its input and return a dictionary with the character frequencies.

Do not use the `collections` module.

<!-- grade_id: cell-7fbaf8610be225a9, points: 0, grade: false, solution: true -->
```python
def char_freq(arg):
    d={}
    uniq=set(arg) # get the distinctive elements in the arg (arg can be a list, a Tuple etc)
    for i in uniq:
        d[i]=0
    #print(d)
    for x in range (0, len(arg)):
        d[arg[x]]=d[arg[x]]+1
    return d
```

<!-- grade_id: cell-ecedb93a6c391539, points: 0, grade: true, solution: false -->
```python
assert char_freq("aba") == {"a": 2, "b": 1}
assert char_freq("") == {}
assert char_freq("ab" * 100 + "cd") == {"a": 100, "b": 100, "c": 1, "d": 1}
```

# 5. Define a function that computes word frequencies in a text.

You can split the text with `text.split()`:

<!-- grade_id: cell-4d4d473dd807f962, points: 0, grade: false, solution: true -->
```python
'''
s=""
t=s.split()
t == []
'''

def word_freq(arg):
    d={}
    splitList=arg.split()
    if splitList == []:
        return {}
    # creat a distinct element list:
    # arg=[1, 2, 0, 1, 1,3,4,4,5,6,6]
    uniq=[]
    uniq.append(splitList[0])
    for x in range (0, len(splitList)):
        for j in range (0,len(uniq)):
            if splitList[x] == uniq[j]:
                break
        else:
            uniq.append(splitList[x])
    #print(uniq)
    for x in range (0,len(uniq)):
        d[uniq[x]]=0
    for x in range (0, len(splitList)):
        d[splitList[x]]= d[splitList[x]]+1
    #print(d)
    return d
#word_freq("the green tea and the black tea")
    #uniq
    

```

<!-- grade_id: cell-1c299c2c935129e7, points: 0, grade: true, solution: false -->
```python
s = "the green tea and the black tea"
assert(word_freq(s) == {"the": 2, "tea": 2, "green": 1, "black": 1, "and": 1})
assert(word_freq("") == {})
```

# Wikipedia module

The following exercises use the `wikipedia` package. The basic usage is illustrated below.

The documentation is available [here](https://pypi.python.org/pypi/wikipedia/).

Searching for pages:

```python
import wikipedia

results = wikipedia.search("Budapest")
results
```

Output:
```
['Budapest',
 'The Grand Budapest Hotel',
 'Budapest Highflyer',
 'Club of Budapest',
 'Budapest Metro',
 'Budapest Short-faced Tumbler',
 'Siege of Budapest',
 'Budapest (disambiguation)',
 'Budapest Ferenc Liszt International Airport',
 'Budapest Gambit']
```

Downloading an article:

```python
article = wikipedia.page("Budapest")
article2 = wikipedia.page("Viet Nam") #+ wikipedia.page("Budapest")


article.summary[:100]
```

Output:
```
'Budapest (, Hungarian pronunciation: [ˈbudɒpɛʃt]) is the capital and the most populous city of Hunga'
```

The content attribute contains the full text:

```python
type(article.content+ article2.content), len(article.content+article2.content)
s=article.content
print(type(s)==str)
```

Output:
```
True
```

By default the module downloads the English Wikipedia. The language can be changed the following way:

```python
wikipedia.set_lang("fr")
```

```python
wikipedia.search("Viet Nam")
```

Output:
```
['Viêt Nam',
 'Guerre du Viêt Nam',
 'République du Viêt Nam',
 'Histoire du Viêt Nam',
 'Têt (Viêt Nam)',
 'Việt Minh',
 'Front national de libération du Sud Viêt Nam',
 'République démocratique du Viêt Nam',
 'Drapeau du Viêt Nam',
 'Économie du Viêt Nam']
```

```python
fr_article = wikipedia.page("Budapest")
fr_article.summary[:100]
```

Output:
```
'Budapest (prononcé [by.da.ˈpɛst] , hongrois : Budapest [ˈbu.dɒ.pɛʃt]  ; allemand : Budapest ou ancie'
```

```python
wikipedia.set_lang("en")
```

# 6. Build a small Wikipedia corpus by downloading pages until their character count exceeds 100000. The corpus should be a list of strings.

You can download arbitrary pages but some pages redirect to disambiguation pages. You should not pick these. Make sure the same page is not included multiple times.

<!-- grade_id: cell-2e7be91b99158dca, points: 0, grade: false, solution: true -->
```python
# Disambifuation error: means the <title> used as the parameter in wikipedia.page(<title>) is the title for many pages (not distinct)
# corpus: a collection of strings
def build_wikipedia_corpus():
    count=0
    s=[]
    searchResult=wikipedia.search("Viet Nam")
    for x in range (0, len(searchResult)):
        article= wikipedia.page(searchResult[x])
        s.append(article.content)
        count = count + len(article.content)
        if count >= 100000:
            break
    return s
s=build_wikipedia_corpus()
type(s[0].split())
```

Output:
```
list
```

<!-- grade_id: cell-f26aa917928c62c8, points: 0, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus()
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 100000

# no duplicate pages
assert len(set(corpus)) == len(corpus)
```

# 7. Add a `min_char` parameter to `build_wikipedia_corpus` that defaults to 10000. You should stop building the corpus when this number is reached.

<!-- grade_id: cell-bf7ffeba0644e753, points: 0, grade: false, solution: true -->
```python
def build_wikipedia_corpus( min_char = 1000):
    count=0
    s=[]
    searchResult=wikipedia.search("Viet Nam")
    for x in range (0, len(searchResult)):
        article= wikipedia.page(searchResult[x])
        s.append(article.content)
        count = count + len(article.content)
        if count >= min_char:
            break
    return s
    
```

<!-- grade_id: cell-68c3de00731699f8, points: 0, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus(1000)
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 1000
# shouldn't be too long
assert sum_length < 100000
```

# 8. Using your previously defined function compute the word frequencies of your Wikipedia corpus. Print the ten most frequent words and their frequency in the following format:

```
a<TAB>12
an<TAB>11
```

You can use str.split() for word tokenization:

```python
"This is a short sentence".split()
l= "" + "ba"
l=="ba"
```

Output:
```
True
```

There are no tests for this task. Try to come up with your own tests.

<!-- grade_id: cell-8c167ceb2a4c65f1, points: 0, grade: true, solution: true -->
```python
# def get_corpus_word_frequencies(corpus):
def get_corpus_word_frequencies(arg=build_wikipedia_corpus()):
    d={} # dictionary to contain word (as key) and its frequency (as value)
    s=arg # a list of strings
    
    # build the collections of words:
    splitStr=[]
    for x in range (0, len(s)):
        splitStr=splitStr+ s[x].split() #s.plit() returns a list
    
    # build a distinct list from the collection:
    uniq=[]
    uniq.append(splitStr[0])
    for x in range (0, len(splitStr)):
        for j in range (0,len(uniq)):
            if splitStr[x] == uniq[j]:
                break
        else:
            uniq.append(splitStr[x])
    
    # count the frequencies of words in the collection:
    for x in range (0,len(uniq)):
        d[uniq[x]]=0
    for x in range (0, len(splitStr)):
        d[splitStr[x]]= d[splitStr[x]]+1
    
    # convert a dictionary to a list: reference:https://www.tutorialspoint.com/How-to-convert-Python-Dictionary-to-a-list
    # for sorting and printing
    l=list(d.items())
    l=sorted(l, key=lambda x: -x[1]) # x represents an element (entry) in the list l, x[1] to access the element of the list entry 
                                     # sort the list in descending order
    for x in range (0, 10):
        print("{0}\t{1}".format(l[x][0],l[x][1]))
    
#build_word_freq()
#t
```

```python
# difference between listA + listB and listA.append(listB)
l=[]
s=[["a","b"], ["c","d"]]
l1=["a"]
l2=["b"]
l=l+l1+l2+s[0]+s[1] 
type(l[0]) # a str


lt=[]
lt.append(s[0])
lt.append(s[1])
type(lt[0]) # a list

d={"c": 2, "b" :1, "a": 3, "d" :0, "f":0}
# convert a dictionary to a list: reference:https://www.tutorialspoint.com/How-to-convert-Python-Dictionary-to-a-list
l=list(d.items())
l
l=sorted(l, key=lambda x: -x[1])
l[0][0] # l2 is a list contain tuples as its element
print("{0} \t {1}".format(l[0][0], l[0][1]))
l
```

Output:
```
a 	 3
[('a', 3), ('c', 2), ('b', 1), ('d', 0), ('f', 0)]
```

Your own tests go here:

<!-- grade_id: cell-d17a9e499c13bd35, points: 0, grade: true, solution: true -->
```python
try:
    get_corpus_word_frequencies()
except:
    assert False
```

Output:
```
the	963
of	568
and	453
in	424
to	260
a	191
Vietnam	175
is	142
The	132
by	127
```

# 9. Define a function that takes a matrix as an input represented as a list of lists. Return its transpose without changing the original matrix.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-adee24575dcbac8e, points: 0, grade: false, solution: true -->
```python
def is_matrix(arg):
    if arg==[]: # if the list is empty
        return True
    len1=len(arg[0])
    for x in range (0, len(arg)):
        if len(arg[x]) != len1:
            return False
    return True

def transpose(arg):
    if is_matrix(arg) == False:
        raise ValueError
        
    # create a tranpose matrix:
    m2=[]
    for x in range (0, len(arg[0])):
        m2.append([])
        
    # add elements into the tranpose matrix:
    for x in range (0, len(m2)):
        for j in range (0, len(arg)):
            m2[x].append(arg[j][x])
    return m2
     
```

<!-- grade_id: cell-9b412cb3d0c0506a, points: 0, grade: true, solution: false -->
```python
m1 = [[1, 2, 3], [4, 5, 6]]
m2 = [[1, 4], [2, 5], [3, 6]]

assert transpose([[1]]) == [[1]]
assert transpose(m1) == m2
# m1 is unchanged
assert m1 == [[1, 2, 3], [4, 5, 6]]

# transpose is an involution (self-inverse)
assert transpose(transpose(m1)) == m1
```

# =================== PASSING LEVEL ===================

# 10. Define a function that takes a matrix as its input and returns the row-wise maximum elements as a list.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-21f91aada57c0332, points: 0, grade: false, solution: true -->
```python

# row-wise maximum elements = the maximum values of each row.

def is_matrix(arg):
    if arg==[]: # if the list is empty
        return True
    len1=len(arg[0])
    for x in range (0, len(arg)):
        if len(arg[x]) != len1:
            return False
    return True

def transpose(arg):
    if is_matrix(arg) == False:
        raise ValueError
        
    # create a tranpose matrix:
    m2=[]
    for x in range (0, len(arg[0])):
        m2.append([])
        
    # add elements into the tranpose matrix:
    for x in range (0, len(m2)):
        for j in range (0, len(arg)):
            m2[x].append(arg[j][x])
    return m2

def rowwise_max( arg):
    new=[]
    if is_matrix(arg) == False:
        raise ValueError
    
    for x in range(0, len(arg)):
        l2=sorted(arg[x]) # sort the list in descending order
        new.append(l2.pop())
    
    #print(new)
    return new


```

```python
# Test cell for myself:
l=[2,1,3,6,5]
l2=sorted(l)
print(l2.pop())
```

Output:
```
6
```

<!-- grade_id: cell-eb440f91ea977032, points: 0, grade: true, solution: false -->
```python
assert rowwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [3, 4, 5]

assert rowwise_max([list(range(1000))]) == [999]
assert rowwise_max(transpose([list(range(1000))])) == list(range(1000))
```

# 11. Define a function that takes a matrix as its input and returns the column-wise maximum elements as a list.

Try to reuse your previous functions.

<!-- grade_id: cell-bde0cb9cf8d75c89, points: 0, grade: false, solution: true -->
```python
def columnwise_max (arg):
    haha=[]
    if is_matrix(arg) == False:
        raise ValueError
    
    new=transpose(arg)
   
    haha=rowwise_max(new)
   
    return haha


```

<!-- grade_id: cell-142e6034ed15005b, points: 0, grade: true, solution: false -->
```python
assert columnwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [1, 5, 3]

assert columnwise_max([list(range(1000))]) == list(range(1000))
assert columnwise_max(transpose([list(range(1000))])) == [999]
```

# 12. Define a function that takes two strings and decides whether they are anagrams. A string is an anagram of another string if its letters can be rearranged so that it equals the other string.¶

For example:

```
abc -- bac
aabb -- abab
```

Counter examples:

```
abc -- aabc
abab -- aaab
```

<!-- grade_id: cell-7ea849b6ffc4421a, points: 0, grade: false, solution: true -->
```python
'''
d1={"a":1, "b":2, "c":3}
d2={"b":2, "a":1, "c":3}
d1==d2 #True
'''

def char_freq(arg): # from exercise 4
    d={}
    uniq=set(arg) # get the distinctive elements in the arg (arg can be a list, a Tuple etc)
    for i in uniq:
        d[i]=0
    #print(d)
    for x in range (0, len(arg)):
        d[arg[x]]=d[arg[x]]+1
    return d

def are_anagrams (arg, arg1):
    d1=char_freq(arg)
    d2=char_freq(arg1)
    if d1==d2:
        return True
    return False


```

<!-- grade_id: cell-231f78d971a5d48b, points: 0, grade: true, solution: false -->
```python
assert(are_anagrams("", "") == True)
assert(are_anagrams("abc", "bac") == True)
assert(are_anagrams("aabb", "abab") == True)
assert(are_anagrams("abab", "aaab") == False)
```

# =================== EXTRA LEVEL ===================

# 13. Define a function that takes a string and returns its bigram frequencies as a dictionary.

Character bigrams are pairs of subsequent characters. For example word apple contains the following bigrams: ap, pp, pl, le.

Do not use the `collections` module.

<!-- grade_id: cell-2669edc3d29b28b6, points: 0, grade: false, solution: true -->
```python
'''
f="the tea"
print(f[7]) # out of range = the string in python doesn't have "\0" (end of string character) like in C/C++
combine=f[0]+f[1]
print(combine) # result: th
print(len(f)) # 7
'''

def get_char_bigrams (arg):
    d={}
    
    # get all the keys for the dictionary:
    for x in range (0,len(arg)-1):        
        combine=arg[x]+arg[x+1]
        d[combine]=0

    
    # count the frequencies:
    for x in range (0,len(arg)-1):
        d[arg[x]+arg[x+1]]=d[arg[x]+arg[x+1]]+1
    
    return d

```

<!-- grade_id: cell-5a3764a1d0909393, points: 0, grade: true, solution: false -->
```python
assert get_char_bigrams("apple") == {'ap': 1, 'pp': 1, 'pl': 1, 'le': 1}
assert get_char_bigrams("apple apple") == {'ap': 2, 'pp': 2, 'pl': 2, 'le': 2, 'e ': 1, ' a': 1}
```

# 14. Define a function that takes two vectors represented as lists and returns their L2 distance.

Also called Euclidean distance.

<!-- grade_id: cell-849b6ea4c65f058e, points: 0, grade: false, solution: true -->
```python
# L2 distance of 2 vectors ex: v1(1,0); v2(3,4) => L2 distance of (v1-v2) = sqrt((1-3)^2 + (0-4)^2)

def l2_distance(arg, arg1):
    if len(arg) != len (arg1):
        print( " two vectors must have equal length")
        raise ValueError
    distanceSquare=0
    for x in range (0, len(arg)):
        distanceSquare=distanceSquare + (arg1[x]-arg[x])**2
        
    distance=distanceSquare**0.5
    return distance
```

<!-- grade_id: cell-4c9f06a8534cbefc, points: 0, grade: true, solution: false -->
```python
assert abs(l2_distance([1, 0], [1, 0]) - 0) < 1e-5
assert abs(l2_distance([1, 0], [0, 1]) - 2 ** 0.5) < 1e-5
```

# 15. Create a function that takes a matrix of size NxM and returns the pairwise L2 distance of each row pair.

For input $A^{N \times M}$, the output should be $D^{N \times N}$ where: $D_{i,j} = L_2 (A_i, A_j)$

<!-- grade_id: cell-644d3391fc995ee9, points: 0, grade: false, solution: true -->
```python
def pairwise_row_distance(arg):
    if is_matrix(arg)==False:
        raise ValueError
    
    D=[]
    for x in range(0, len(arg)):
       D.append([])
    
    for i in range (0, len(arg)):
        for j in range (0, len(arg)):
            D[i].append(l2_distance(arg[i],arg[j]))
    
    return D

```

<!-- grade_id: cell-8bc3cf03b303c7c3, points: 0, grade: true, solution: false -->
```python
D = pairwise_row_distance([[1, 0, 2], [0, 2, 1]])
# should be a matrix
assert is_matrix(D) == True
assert len(D) == 2
assert len(D[0]) == 2
assert D[0][0] == 0
assert abs(D[0][1] - 2.449489742783178) < 1e-3
assert abs(D[1][0] - 2.449489742783178) < 1e-3
assert D[1][1] == 0

# should be symmetric
assert transpose(D) == D
```